Atypical Biological Motion Perception in Children with Attention Deficit Hyperactivity Disorder: Local Motion and Global Configuration Processing

Author response:

The following is the authors’ response to the previous reviews.

eLife assessment

This manuscript represents a cleanly designed experiment for assessing biological motion processing in children (mean age = 9) with and without ADHD. The group differences concerning accuracy in global and local motion processing abilities are solid, but the analyses suggesting dissociable relationships between global and local processing and social skills, age, and IQ are inconclusive. The results are useful in terms of understanding ADHD and the ontogenesis of different components of the processing of biological motion.

We thank the editors and reviewers for their valuable feedback and constructive comments. We have carefully considered each point raised by the reviewers and made the necessary revisions to the manuscript. Regarding the relationships between global and local BM processing, the accumulated evidence from previous studies has converged on the dissociation of the two BM components, e.g., while global BM processing is susceptible to learning and practice, local BM processing does not show a learning trend (Chang and Troje, 2009; Grossman et al., 2004), and the brain activations in response to local and global BM cues are different (Chang et al., 2018; Duarte et al., 2022). Nevertheless, we concurred with reviewers that the evidence for such dissociation from the current study by itself is not strong enough. Therefore, we have toned down on this point and no longer claimed the dissociation (including the title). Based on the current results, we focused our discussion on the different aspects of BM processing in children with and without ADHD.

Public Reviews:

Reviewer #1 (Public Review):

Summary:

The paper presents a nice study investigating the impairments of biological motion perception in individuals with ADHD in comparison with neurotypical controls. Motivated by the idea that there is a relationship between biological motion perception and social capabilities, the authors investigated the impairments of local and global (holistic) biological motion perception, the diagnosis status, and several additional behavioral variables that are affected in ADHS (IQ, social responsiveness, and attention / impulsivity). As well local as global biological motion perception is impaired in ADHD individuals. In addition, the study demonstrates a significant correlation between local biological motion perception skills and the social responsiveness score in the ADHD group, but not in controls. A path analysis in the ADHD group suggests that general performance in biological motion perception is influenced mainly by global biological motion perception performance and attentional and perceptual reasoning skills.

Strengths:

It is true that there exists not much work on biological motion perception and ADHD. Therefore, the presented study contributes an interesting new result to the biological motion literature, and adds potentially also new behavioral markers for this clinical group. The design of the study is straightforward and technically sound, and the drawn conclusions are supported by the presented results.

Thanks for this positive assessment of our work.

Weaknesses:

Some of the claims about the relationship between genetic factors and ADHD and the components of biological motion processing have to remain speculative at this point because genetic influences were not explicitly tested in this paper. Specifically, the hypothesis that the perception of human social interaction is critically based on a local mechanism for the detection of asymmetry in foot trajectories of walkers (this is what 'BL-local' really measures), or on the detection of live agents in cluttered scenes seems not very plausible.

Thanks for these comments. We agree that the relationship between genetic factors and BM perception remains to be further examined, as we did not test the genetic influences in this study. We have deleted relavant discussion about genetics. Based on our results, we discuss the possible mechanisms behind the relationship between local BM processing and social interaction in the revised manuscript as follows:

“As mentioned above, we found a significant negative correlation between the SRS total score and the accuracy of local BM processing, specifically in the ADHD group. This could be due to decreased visual input related to atypical local BM processing, which further impairs global BM processing. According to the two-process theory of biological motion processing61, local BM cues guide visual attention towards BM stimuli55,62. Consequently, the visual input of BM stimuli increases, facilitating the development of the ability to process global BM cues through learning21,63. The latter is a prerequisite for attributing intentions to others and facilitating social interactions with other individuals20,64,65. Thus, atypical local BM processing may contribute to impaired social interaction through altered visual inputs. Further empirical studies are required to confirm these hypotheses.” (lines 417 - 428)

Based on my last comments, now the discussion has been changed in a way that tries to justify the speculative claims by citing a lot of other speculative papers, which does not really address the problem. For example, the fact that chicks walk towards biological motion stimuli is interesting. To derive that this verifies a fundamental mechanism in human biological motion processing is extremely questionable, given that birds do not even have a cortex. Taking the argumentation of the authors serious, one would have to assume that the 'Local BM' mechanism is probably located in the mesencephalon in humans, and then would have to interact in some way with social perception differences of ADHD children. To me all this seems to make very strong (over-)claims. I suggest providing a much more modest interpretation of the interesting experimental result, based on what has been really experimentally shown by the authors and closely related other data, rather than providing lots of far-reaching speculations.

In the same direction, in my view, go claims like 'local BM is an intrinsic trait' (L. 448) , which is not only imprecise (maybe better 'mechanisms of processing of local BM cues') but also rather questionable. Likely, this' local processing of BM' is a lower level mechanisms, located probably in early and mid-levels of the visual cortex, with a possible influence of lower structures. It seems not really plausible that this is related to a classical trait variables in the sense of psychology, like personality, as seems to be suggested here. Also here I suggest a much more moderate and less speculative interpretation of the results.

We thank the reviewer for pointing out these issues. According to these comments, we have carefully revised the discussion to avoid strong (over-) claims. We have deleted the example of chicks, but substituted with more empirical studies to explain our results. We agree that the Local BM mechanism is probably located in subcortical regions in humans, which were reported by some MRI studies (Chang et al., 2018; Hirai and Senju, 2020; Loula et al., 2005). We have added some evidence that atypical local BM processing may decrease visual inputs related to social information as follows:

“According to the two-process theory of biological motion processing61, local BM cues guide visual attention towards BM stimuli55,62. Consequently, the visual input of BM stimuli increases, facilitating the development of the ability to process global BM cues through learning21,63. The latter is a prerequisite for attributing intentions to others and facilitating social interactions with other individuals20,64,65. Thus, atypical local BM processing may contribute to impaired social interaction through altered visual inputs.” (lines 421 - 427)

We have also deleted the clarims of 'local BM is an intrinsic trait' (originally L. 448) and related discussion as it was not conclusive based on the current study.

Reviewer #2 (Public Review):

Summary:

Tian et al. aimed to assess differences in biological motion (BM) perception between children with and without ADHD, as well as relationships to indices of social functioning and possible predictors of BM perception (including demographics, reasoning ability and inattention). In their study, children with ADHD showed poorer performance relative to typically developing children in three tasks measuring local, global, and general BM perception. The authors further observed that across the whole sample, performance in all three BM tasks was negatively correlated with scores on the social responsiveness scale (SRS), whereas within groups a significant relationship to SRS scores was only observed in the ADHD group and for the local BM task. Local and global BM perception showed a dissociation in that global BM processing was predicted by age, while local BM perception was not. Finally, general (local & global combined) BM processing was predicted by age and global BM processing, while reasoning ability mediated the effect of inattention on BM processing.

Strengths:

Overall, the manuscript is presented in a clear fashion and methods and materials are presented with sufficient detail so the study could be reproduced by independent researchers. The study uses an innovative, albeit not novel, paradigm to investigate two independent processes underlying BM perception. The results are novel and have the potential to have wide-reaching impact on multiple fields.

We appreciate the reviewer’s positive feedback very much.

Weaknesses:

The manuscript has greatly improved in clarity and methodological considerations in response to the review. There are only a few minor points which deserve the authors' attention:

When outlining the moviation for the current study, results from studies in ADHD and ASD are used too interchangeably. The authors use a lack of evidence for contributing (psychological/developmental) factors on BM processing in ASD to motivate the present study and refer to evidence for differences between typical and non-typical BM processing using studies in both ASD and ADHD. While there are certainly overlapping features between the two conditions/neurotypes, they are not to be considered identical and may have distinct etiologies, therefore the distinction between the two should be made clearer.

We thank the reviewer for pointing out this issue. We have removed some unnecessary citations about ASD and referred to studies about social cognition in ADHD to elaborate the motivation of this study:

“Further exploration of a diverse range of social cognitions (e.g., biological motion perception) can provide a fresh perspective on the impaired social function observed in ADHD. Moreover, recent studies have indicated that the social cognition in ADHD may vary depending on different factors at the cognitive, pathological, or developmental levels, such as general cognitive impairment5, symptoms severity8, or age5. Nevertheless, understanding how these factors relate to social cognitive dysfunction of in ADHD is still in its infancy. Bridging this gap is crucial as it can help depict the developmental trajectory of social cognition and identify effective interventions for impaired social interaction in individuals with ADHD.” (lines 53 - 62)

In the first/main analysis, is unclear to me why in the revised manuscript the authors changed the statistical method from ANOVA/ANCOVA to independent samples t-tests (unless the latter were only used for post-hoc comparisons, then this needs to be stated). Furthermore, although p-values look robust, for this analysis too it should be indicated whether and how multiple comparison problems were accounted for.

Thanks for the reviewer’s comments. According to the suggestions from reviewer #3, it may be inapposite to regard gender as a covariate in ANOVA, which may violate the assumptions of ANCOVA. To ensure that gender does not influence the results, firstly, we separated boys and girls on the plots with different coloured individual data points, and there are no signs of a gender effect in their TD group. Secondly, we use t-tests to examine the difference between TD and ADHD groups. Finally, we conducted a subsampling analysis with balanced data, and the results remained consistent.

In part 1 of the results, we aimed to compare the task accuracies between the TD and ADHD groups in three independent tasks, which assess the participants’ abilities to process three types of BM cues. We assumed that individuals with ADHD show poorer performance in three tasks compared to TD individuals. With regard to that, we consider that multiple comparisons may not be necessary.

Reviewer #3 (Public Review):

Strengths:

The authors present differences between ADHD and TD children in biological motion processing, and this question has not received as much attention as equivalent processing capabilities in autism. They use a task that appears well controlled. They raise some interesting mechanistic possibilities for differences in local and global motion processing, which are distinctions worth exploring. The group differences will therefore be of interest to those studying ADHD, as well as other developmental conditions, and those examining biological motion processing mechanisms in general.

We appreciate the reviewer’s positive assessment of this work.

Weaknesses:

The data are not strong enough to support claims about differences between global and lobal processing wrt social communication skills and age. The mechanistic possibilities for why these abilities may dissociate in such a way are interesting, but the crucial tests of differences between correlations do not present a clear picture. Further empirical work would be needed to test the authors' claims. Specifics:

The authors state frequently that it was the local BM task that related to social communication skills (SRS) and not the global tasks. However, the results section shows a correlation between SRS and all three tasks. The only difference is that when looking specifically within the ADHD group, the correlation is only significant for the local task. The supplementary materials demonstrate that tests of differences between correlations present an incomplete picture. Currently they have small samples for correlations, so this is unsurprising.

Thanks for this comment. We agree with the reviewer that the relationship between local and global processing with social communication and age needs more expirical work. Based on our results, there are only possible dissociable roles of local and global BM processing. The accumulated evidence from previous studies has converged on this dissociation, e.g., whild global BM processing is susceptible to learning and practice, local BM processing does not show a learning trend (Chang and Troje, 2009; Grossman et al., 2004), and the brain activations in response to local and global BM cues are different (Chang et al., 2018; Duarte et al., 2022). We concurred with reviewers that the evidence for such dissociation from the current study by itself is not strong enough. Therefore, we have toned down on this point and no longer emphasized the dissociation. Based on the current results, we focused our discussion on the different aspects of BM processing in children with and without ADHD. Future studies with larger sample sizes are needed to confirm this disociable relationship.

Theoretical assumptions. The authors make some statements about local vs global biological motion processing that should still be made more tentatively. They assume that local processing is specifically genetically whereas global processing is a product of experience. These data in newborn chicks are controversial and confounded - I cannot remember the specifics but I think there an upper vs lower visual field complexity difference here.

We appreciate the reviewer’s suggestion. We agree that the relationship between genetic factors and BM perception remains to be further examined as we didn’t perform any genetic analysis in the current study. Some speculative papers have been removed, so do the statement about newborn chicks given the controversial and confounded results. We have ﻿toned down our claims and povided a moderate interpretation of the results:

“Sensitivity to local BM cues emerges early in life54,55 and involves rapid processing in the subcortical regions16,56-58. As a basic pre-attentive feature23, local BM cues can guide visual attention spontaneously59,60. In contrary, the ability to process global BM cues is related to slow cortical BM processing and is influenced by many factors such as attention25,26 and visual experience21,51. As mentioned above, we found a significant negative correlation between the SRS total score and the accuracy of local BM processing, specifically in the ADHD group. This could be due to decreased visual input related to atypical local BM processing, which further impairs global BM processing. According to the two-process theory of biological motion processing61, local BM cues guide visual attention towards BM stimuli55,62. Consequently, the visual input of BM stimuli increases, facilitating the development of the ability to process global BM cues through learning21,63. The latter is a prerequisite for attributing intentions to others and facilitating social interactions with other individuals20,64,65. Thus, atypical local BM processing may contribute to impaired social interaction through altered visual inputs.” (lines 413 - 427)

“Few developmental studies have been conducted on local BM processing. The ability to process local BM cues remained stable and did not exhibit a learning trend21,25. A reasonable interpretation may be that local BM processing is a low-level mechanism, probably performed by the primary visual cortex and subcortical regions such as the superior colliculus, pulvinar, and ventral lateral nucleus14,56,61.” (lines 441- 446)

Readability. The manuscript needs very careful proofreading and correction for grammar. There are grammatical errors throughout.

Thank the reviewer for this feedback. We have performed thorough proofreading and corrected grammatical errors throughout the manuscript.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

I thank the authors for their revisions that address several of the minor points that I raised in my last review. A number of requests are still not sufficiently answered:

L. 290 ff.: These model 'BM-local = age + gender etc ' is a pretty sloppy notation. I think what is meant that a GLM was used that uses the predictors genderetc. time appropriate beta_i values. This formulas should be corrected or one just says that a GLM was run with the predictors gender

The same criticism applies to these other models that follow.

This was corrected.

However, the corrected text remains sloppy: example: 'BM-locaL = ...' What exacty is 'BM-Local' the accuracy? etc. Here a precise notation shoudl be given that clearly names which variables are used here as predictors and target variables.

We appreciate the reviewer’s suggestion. We clarified which variables are used in our model and gived them precise notations:

“Three linear models were built to investigate the contributing factors: (a) ACClocal = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention, (b) ACCglobal = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention, and (c) ACCgeneral = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention + β5 * ACClocal + β6 * ACCglobal. ACClocal, ACCglobal and ACCgeneral refer to the response accuracies of the three tasks in the ADHD group, and QbInattention is the standardised score for sustained attention function.” (lines 337 - 343)

All these models assume linearity of the combination of the predictors. was this assumption verified?

We referred to the previous study of BM perception in children. They found main predictor variables, including IQ (Rutherford et al., 2012; Jones et al., 2011) and age (Annaz et al., 2010; van et al., 2016), have a linear relation with the ability of BM processing.

This answer is insufficient and not convincing. Because a variable Y depends linearly on predictor A and B in some other study, this does not imply that is is also linear in predictor C, or does not show interactions with such predictors in the present study.

What is needed here is the testing of models with interaction terms and verifying that such models are not better predictors. If authors do not want to do this, they need at least to clearly point out that they made the strong assumption of linearity of their model, which might be wrong and thus be a substantial limitation of their analysis.

Thanks for the suggestion. We tried to compare each possible mode with and without relative interactions. The results showed that the change of Coefficient of Determination (R-squared, R2) between the two models was not statistically significant.

L. 296ff.: For model (b) it looks like general BM performance is strongly driven by the predictor global BM performance in the ADHD group. Does the same observation also apply to the controls?

The same phenomenon was not observed in TD children. We have briefly discussed this point in the Discussion section of the revised manuscript (lines 449 - 459).

Was such a path analysis also done for the TD subjects or not? If yes, was then also predicted that the variable BM-Global largely and directedly influences the variable BM-General? (The answer refers to the general discussion section, where no such analysis is presented, as far as I understand.)

Thank you for your comment. We also conduct a path analysis similar to that in the ADHD group. There is no statistically significant mediator effect in the TD group. Please see Figure S3 for complete statistics.

Reviewer #2 (Recommendations For The Authors):

(1) Please add public access to the data repository so data availability can be assessed.

The data analyzed during the study is available at https://osf.io/37p5s/.

(2) Lines 119-115: The differences observed in ADHD participants in the studies referenced here were relative to what group? The last sentence here also refers to two groups, and it is difficult to gather which specific groups are meant, also because the two references relate to both ADHD and ASD samples. Please clarify.

The suggestion is well taken. We have clarified the expressions accordingly:

“Specifically, compared with the typically developing (TD) group, children with ADHD showed reduced activity of motion-sensitive components (N200) while watching biological and scrambled motions, although no behavioural differences were observed. Another study found that children with ADHD performed worse in BM detection with moderate noise ratios than the TD group32.” (lines 100 - 105)

(3) Line 116: I'm not sure what is meant by 'despite initial indications' - please briefly specify/summarise here why the investigation into BM processing in ADHD is warranted.

Thank the reviewer for pointing out this issue. We rephrase this part and briefly specify “why the investigation into BM processing in ADHD is warranted”:

“Despite initial findings about atypical BM perception in ADHD, previous studies on ADHD treated BM perception as a single entity, which may have led to misleading or inconsistent findings28. Hence, it is essential to deconstruct BM processing into multiple components and motion features.” (lines 108 -111)

(4) Lines 290-293: Please complete the sentence.

Thank the reviewer for pointing out this issue. Th sentence has been completed:

“For Task 2 and 3, where children were asked to detect the presence or discriminate the facing direction of the target walker, TD group have higher accuracies than the ADHD group (Task 2 - TD: 0.70 ± 0.12, ADHD: 0.59 ± 0.12, t73 = 3.677, p < 0.001, Cohen's d = 0.861; Task 3 - TD: 0.79 ± 0.12, ADHD: 0.63 ± 0.17, t73 = 4.702, p < 0.001, Cohen's d = 1.100).” (lines 284 - 288)

Reviewer #3 (Recommendations For The Authors):

(1) Conclusions concerning differences between the local and global tasks wrt SRS and age (see above). I believe the authors need to reword throughout to reflect that the tests of differences between these crucial correlations did not present a clear picture.

We have reworded throughout the paper to reflect the inconclusiveness with regard to the relationship between local and global processing with social communication based on this study only. Future studies with larger sample sizes are needed to confirm this conclusion. The mechanism for this dissociable relationship should be validated by more psychologial tests in the future studies.

(2) I would again tone down the discussion of genetic specification of local processing, given it is highly controversial.

We thank the reviewer for pointing out the issue. We agree the point about the genetic specification of local processing remains controversial. The interpretation of results about local BM processing has been rephrased. Please refer to our response to the point #2 mentioned.

(3) The manuscript needs very careful proofreading and grammatical correction throughout.

Thanks for the suggestion to check the grammar. We have carefully proofread the manuscript to correct grammatical errors

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The ability to perceive biological motion (BM) is crucial for human survival and social interaction. Plentiful studies have found impaired BM perception in autism spectrum disorder characterized by deficits in social interaction. Children with attention deficit hyperactivity disorder (ADHD) often exhibit similar deficits in social interaction, but few studies have investigated BM perception in ADHD. Here, we compared the differences in abilities to process local kinematic and global configurational cues, two fundamental abilities of BM perception, between typically developing (TD) and ADHD children. Then, we further investigated the relationship between BM perception and social interaction skills measured by the Social Responsiveness Scale (SRS) and examined the contributions of latent factors (e.g., gender, age, attention and intelligence) to BM perception. Results revealed that children with ADHD exhibited atypical BM perception with a potential dissociation between local and global BM information processing. Local BM processing ability was largely related to social interaction skills, whereas global BM processing ability would significantly improve with age. Critically, general BM perception (i.e., both local and global BM cues) could be affected by sustained attention ability in children with ADHD. This relationship was mainly mediated by Reasoning Intelligence. These findings elucidate the atypical biological motion perception in ADHD and the latent factors related to BM perception. Moreover, this study provides new evidence for local BM perception as a hallmark of social cognition and advances the comprehensive understanding of the distinct roles of local and global processing in BM perception and social cognitive disorders.

Revised version: Perceiving biological motion (BM) is crucial for human survival and social interaction. Many studies have reported impaired BM perception in autism spectrum disorder, which is characterised by deficits in social interaction. Children with attention deficit hyperactivity disorder (ADHD) often exhibit similar difficulties in social interaction. However, few studies have investigated BM perception in children with ADHD. Here, we compared differences in the ability to process local kinematic and global configurational cues, two fundamental abilities of BM perception, between typically developing and ADHD children. We further investigated the relationship between BM perception and social interaction skills measured using the Social Responsiveness Scale and examined the contributions of latent factors (e.g., sex, age, attention, and intelligence) to BM perception. The results revealed that children with ADHD exhibited atypical BM perception. Local and global BM processing showed distinct features. Local BM processing ability was related to social interaction skills, whereas global BM processing ability significantly improved with age. Critically, general BM perception (i.e., both local and global BM processing) may be affected by sustained attentional ability in children with ADHD. This relationship was primarily mediated by reasoning intelligence. These findings elucidate atypical BM perception in ADHD and the latent factors related to BM perception. Moreover, this study provides new evidence that BM perception is a hallmark of social cognition and advances our understanding of the potential roles of local and global processing in BM perception and social cognitive disorders.

Introduction

Earlier version: Attention deficit/hyperactivity disorder (ADHD) is a common developmental disorder, with a prevalence ranging from 2% to 7% in children and adolescents, averaging around 5%1. Besides the well-established core symptoms of ADHD (including a deficit of sustained attention, hyperactivity, and impulsivity), some autism spectrum disorder (ASD) characteristics, such as dysfunctions in social communication and social interaction, have also been frequently reported in ADHD children2–4. However, experimental studies of social cognition in ADHD are still scarce. Some studies described worse performance in tasks of social cognition. Among them, impaired theory of mind (ToM) and emotion recognition are most frequently reported5–7. It is difficult for children with ADHD to recognize others’ emotions and intentions. However, the understanding of other social cognition processes in ADHD remains limited. Further exploring a diverse range of social cognition (e.g., biological motion perception) can provide a fresh perspective on the impaired social function observed in ADHD. Moreover, recent studies indicated that the social cognition in ADHD may vary depending on different factors at the cognitive, pathological, or developmental levels, such as general cognitive impairment5, the severity of symptoms8, or age5. Yet, how these factors relate to the dysfunctions of social cognition in ADHD is still in its infancy. Bridging the gap is important as it can help depict the developmental trajectory of social cognition and identify effective interventions for impaired social interaction in individuals with ADHD.

Revised version: Attention deficit hyperactivity disorder (ADHD) is a common developmental disorder with a prevalence ranging from 2% to 7% in children and adolescents, averaging approximately 5%1. In addition to the well-established core symptoms of ADHD (including inability to sustained attention, hyperactivity, and impulsivity), some characteristics of autism spectrum disorder (ASD), such as dysfunctions in social communication and social interaction, have also been frequently observed in children with ADHD2–4. Nevertheless, experimental studies focusing on social cognition in children with ADHD are limited. Some studies have reported poor performance on social cognition tasks. Among these, impaired theory of mind (ToM) and emotion recognition are the most frequently reported5–7. It is difficult for children with ADHD to recognise others’ emotions and intentions of others. However, our understanding of other social cognitive processes in ADHD remains limited. Further exploration of a diverse range of social cognitions (e.g., biological motion perception) can provide a fresh perspective on the impaired social function observed in ADHD. Moreover, recent studies have indicated that the social cognition in ADHD may vary depending on different factors at the cognitive, pathological, or developmental levels, such as general cognitive impairment5, symptoms severity8, or age5. Nevertheless, understanding how these factors relate to social cognitive dysfunction of in ADHD is still in its infancy. Bridging this gap is crucial as it can help depict the developmental trajectory of social cognition and identify effective interventions for impaired social interaction in individuals with ADHD.

Earlier version: Biological motion (BM), which refers to the movements of a living creature, conveys a wealth of information beyond bodily movements9, such as intention10, emotion11, gender12, and identity13,14. The advent of point-light displays (PLD) technology used to depict human motions15 allowed researchers to separate biological motion from other characteristics like shape and color. Considering the seminal impact on cognitive science, developmental and clinical neuroscience, BM perception has drawn significant attention from scientists. Some researchers attempted to deconstruct BM processing into more specific components. Our study concentrated on two fundamental abilities involved in processing BM cues (Figure 1): the ability to process local BM cues derived from major joint motion tracks and the ability to process global BM cues of human configuration. Previous studies provided some findings about the difference between local and global BM perception. Separable neural signals for these abilities implied two independent BM processing stages16–18. The ability to process local BM cues appears to be genetically highly preserved19. For example, newly hatched chicks without visual experience exhibit a spontaneous preference for BM20. A similar finding is also reported in 2-day-old babies21. Local BM cues not only carry locomotive direction22 but also contribute to detecting life in the visual environment23 without observers’ explicit recognition or attention24–27. As a result, the processing of local BM cues is unaffected by attention, relatively robust to masking noise and does not show a learning trend27,28. In contrast, global BM processing involves top-down modulation, with attention playing a critical role in its perception28,29. Dispersed attention adversely affects its performance. Compared with local BM processing, global BM processing is susceptible to learning and is heavily hindered by increasing mask densities28. These findings suggest that local and global mechanisms might play different roles in BM perception, though the exact mechanisms underlying the distinction remain unclear. Exploring the two components of BM perception will enhance our understanding of the difference between local and global BM processing, shedding light on the psychological processes involved in atypical BM perception.

Revised version: Biological motion (BM), which refers to the movement of a living creature, conveys a wealth of information beyond bodily movements9, such as intention10, emotion11, sex12, and identity13,14. The advent of point-light displays (PLD) technology, which is used to depict human motions,15 allows researchers to separate biological motion from other characteristics such as shape and colour. Considering its seminal impact on cognitive, developmental, and clinical neuroscience, BM perception has drawn significant attention from scientists. Some researchers have attempted to deconstruct BM processing into more specific components. Our study concentrated on two fundamental abilities involved in processing BM cues (Figure 1): the ability to process local BM cues derived from major joint motion tracks, and the ability to process global BM cues of human configuration. Previous studies revealed differences between local and global BM perception. Separate neural signals for these abilities imply two independent BM processing stages16–18. Local BM cues not only help identify locomotive direction19 but also contribute to the detecting of life in the visual environment20 without the observers’ explicit recognition or attention21–24. As a result, the processing of local BM cues is unaffected by attention, relatively robust to masking noise, and does not show a learning trend24,25. In contrast, global BM processing involves top-down modulation, with attention playing a critical role in its perception25,26. Dispersed attention adversely affects performance. Compared with local BM processing, global BM processing is susceptible to learning and is heavily hindered by increased mask densities25. These findings suggest that local and global mechanisms play different roles in BM perception, although the exact mechanisms underlying this distinction remains unclear. Exploring these two components of BM perception will enhance our understanding of the differences between local and global BM processing and shed light on the psychological processes involved in atypical BM perception.

Earlier version: Schematic representation of Biological Motion (BM) and scrambled BM sequence. Intact walker contains information of local kinematics and global configuration. Local kinematics refers to the motion tracks of each critical joint illustrated by the chromatic dot. Global configuration is composed of the relative locations of each joint. In scrambled BM sequence, global configuration cues have been deprived, but local kinematics have been retained. (Figure reconstructed from Wang et al., 201819)

Revised version: Schematic representation of Biological Motion (BM) and scrambled BM sequence. Intact walker contains information of local kinematics and global configuration. Local kinematics refers to the motion tracks of each critical joint illustrated by the chromatic dot. Global configuration is composed of the relative locations of each joint. In the scrambled BM sequence, global configuration cues have been removed, but local kinematics have been retained. (Figure reconstructed from Wang et al., 201830)

Removed in the revised version: In recent years, BM perception has received significant attention in studies of mental disorders (e.g., schizophrenia30) and developmental disabilities, particularly in ASD, characterized by deficits in social communication and social interaction31,32. An important reason is that BM perception is considered a hallmark of social cognition. Individuals with deficits in BM processing exhibit worse social perception in daily life10. The other study found that participants’ ability to process BM cues was correlated with their autistic traits, especially in the subdimension of social communication19. Notably, the contributing factors of impaired BM perception in children with ASD is unclear. Previous meta-analysis studies of BM perception in ASD found mixed results about the effects of age and IQ on processing BM31,32.

Earlier version: Therefore, it is essential to focus on revealing differences in BM processing between typical and atypical development and identified potential factors that may influence BM perception, which will enhance our comprehension of social dysfunction in atypical development. Compared with numerous studies examining impaired BM perception in ASD, few studies focused on BM perception in children with ADHD. An EEG study found neuro-electrophysiological changes in processing BM stimuli in ADHD33. Specifically, children with ADHD showed reduced activity in motion- sensitive components (N200) while watching biological and scrambled motion, although no behavioral difference was observed. Another study found that children with ADHD performed worse in BM detection with moderate ratios of noise34. This may be due to the fact that BM stimuli with noise dots will increase the difficulty of identification, which highlights the difference in processing BM between the two groups33,35.

Revised version: In recent years, BM perception has received significant attention in studies on mental disorders (e.g., schizophrenia27) and developmental disabilities, particularly in ASD, which is characterised by deficits in social communication and social interaction28,29. This is because BM perception is considered a hallmark of social cognition. Individuals with deficits in BM processing exhibit worse social perception in daily life10. Another study found that participants’ ability to process BM cues correlated with their autistic traits, particularly in the subdimension of social communication30. Therefore, examining BM perception could enhance our understanding of social dysfunction in children with ADHD. Compared with the numerous studies examining impaired BM perception in ASD, few studies have focused on BM perception in children with ADHD. An EEG study found neuroelectrophysiological changes in the processing of BM stimuli in children with ADHD31. Specifically, compared with the typically developing (TD) group, children with ADHD showed reduced activity of motion-sensitive components (N200) while watching biological and scrambled motions, although no behavioural differences were observed. Another study found that children with ADHD performed worse in BM detection with moderate noise ratios than the TD group32. This finding may be due to the fact that BM stimuli with noise dots will increase the difficulty of identification33, which highlights the difference in BM processing between TD and ADHD groups31.

Earlier version: Despite initial indications, a comprehensive investigation into BM perception in ADHD is warranted. We proposed that it is essential to deconstruct BM processing into its multiple components and motion features, since treating them as a single entity may lead to misleading or inconsistent findings31. To address this issue, we employed a carefully-designed behavioral paradigm used in our previous study19, making slight adjustments to adapt for children. This paradigm comprises three tasks. Task 1 (BM-local) aimed to assess the ability to process local BM cues. Scrambled BM sequences were displayed and participants could use local BM cues to judge the facing direction of the scrambled walker. Task 2 (BM-global) tested the ability to process the global configuration cues of the BM walker. Local cues were uninformative, and participants used global BM cues to determine the presence of an intact walker. Task 3 (BM-general) tested the ability to process general BM cues (local + global cues). The stimulus sequences consisted of an intact walker and a mask containing similar target local cues, so participants could use general BM cues (local + global cues) to judge the facing direction of the walker.

Revised version: Despite initial findings about atypical BM perception in ADHD, previous studies on ADHD treated BM perception as a single entity, which may have led to misleading or inconsistent findings28. Hence, it is essential to deconstruct BM processing into multiple components and motion features. To enhance our understanding of the ability to process distinct BM cues in ADHD, we employed a carefully designed behavioural paradigm, as used in our previous study30, with slight adjustments made to adapt it for children. This paradigm comprised three tasks: Task 1 (BM-local), Task 2 (BM-global), and Task 3 (BM-general). Task 1 (BM-local) assessed the ability to process local BM cues. Scrambled BM sequences were displayed and the participants used local BM cues to judge the facing direction of the scrambled walker. Task 2 (BM-global) tested the ability to process the global configuration cues of the BM walker. Local cues were uninformative, and the participants used global BM cues to determine the presence of an intact walker. Task 3 (BM-general) tested participants’ ability to process general BM cues (local and global cues). The stimulus sequences consisted of an intact walker and a mask containing similar target local cues so that the participants could use general BM cues (local + global cues) to judge the direction the walker was facing.

Earlier version: In Experiment 1, we examined three specific BM perception abilities in children with ADHD. As mentioned earlier, children with ADHD also show impaired social interaction2–4, which implies atypical social cognition. Therefore, we speculated that children with ADHD performed worse in the three tasks compared to TD children. In Experiment 2, we further explored the relationship between BM perception and social interaction ability in ADHD, as well as identified potential factors (e.g., IQ, age and attention) that may affect BM perception in this population. We speculate that if local motion and global configuration of BM are indeed contributed by distinct mechanisms as suggested by previous studies, then their impairments in the ADHD population and the influential factors behind the impairments may be different.

Revised version: Experiment 1 examined three specific BM perception abilities in children with ADHD. Children with ADHD show impaired social interaction2–4, which implies atypical social cognition. Therefore, we speculated that children with ADHD would performed worse on the three tasks than TD children. In Experiment 2, we further explored the relationship between BM perception and social interaction ability in children with ADHD and identified potential factors (e.g., intelligence quotient [IQ], age, and attention) that may affect BM perception in this population. We speculated that if the mechanisms of processing local and global BM cues are indeed distinct, as suggested by previous studies, then impairments in the ADHD population and the influential factors behind the impairment may be different.

Materials and methods

Participants

Earlier version: We recruited a total of one hundred and seventeen children with and without ADHD to enter this study. Eighty-one children met the ADHD diagnostic criteria of the Diagnostic and Statistical Manual of Mental Disorders (DSM-5)36. The clinical diagnosis was first made by an experienced child and adolescent psychiatrist in the Child and Adolescent Psychiatric Outpatient Department of the Peking University Sixth Hospital, based on the ADHD Rating Scale. The Chinese version of the kiddie- schedule for affective disorders and schizophrenia-present and lifetime version DSM- 5 (K-SADS-PL-C DSM-5)37,38, a semi-structured interview instrument, was then implemented to confirm the diagnosis. Thirty-six typically developing (TD) children were from ordinary primary schools in Beijing and screened out the presence of ADHD, ASD, affective disorders and behavioral disorders by a trained psychiatrist. All subjects of ADHD groups have full-scale IQ above 75 (5th upper percentile) on Wechsler Intelligence Scale for Children-Fourth Edition, and all TD children had a full-scale IQ that fell above the 5th percentile on the Raven’s Standard Progressive Matrices39 which is used in measuring reasoning ability and regarded as a non-verbal estimate of intelligence. Exclusion criteria for both groups were: (a) neurological diseases; (b) other neurodevelopmental disorders (e.g., ASD, Mental retardation, and tic disorders), affective disorders and schizophrenia; (c) that would impact experiments completion; (d) taking psychotropic drugs or stimulants in 30 days. (e) previous head trauma or neurosurgery.

Revised version: One hundred and seventeen children with and without ADHD were recruited for this study. Eighty-one children met the ADHD diagnostic criteria of the Diagnostic and Statistical Manual of Mental Disorders (DSM-5)34. The clinical diagnosis was first made by an experienced child and adolescent psychiatrist in the Child and Adolescent Psychiatric Outpatient Department of Peking University Sixth Hospital, based on the ADHD Rating Scale. The Chinese version of the Kiddie-Schedule for Affective Disorders and Schizophrenia-Present and Lifetime Version DSM-5 (K-SADS-PL-C DSM-5)35,36, a semi-structured interview instrument, was then implemented to confirm the diagnosis. Thirty-six TD children were from ordinary primary schools in Beijing were screened for the presence of ADHD, ASD, affective disorders, and behavioural disorders by a trained psychiatrist. All participants in the ADHD groups had a full-scale IQ > 75 (5th upper percentile) on the Wechsler Intelligence Scale for Children-Fourth Edition, and all TD children had a full-scale IQ above the 5th percentile on the Raven’s Standard Progressive Matrices37, which is used to measure reasoning ability and is regarded as a non-verbal estimate of intelligence. The exclusion criteria for both groups were as follows: (a) neurological diseases; (b) other neurodevelopmental disorders (e.g., ASD, mental retardation, and tic disorders), affective disorders, and schizophrenia; (c) disorders that would impact the completion of the experiment; (d) taking psychotropic drugs or stimulants within the past 30 days; and (e) previous head trauma or neurosurgery.

Earlier version: Thirty-six TD children (age = 9.09 ± 2.18, 14 male) and thirty-nine children with ADHD (age = 9.88 ± 2.23, 28 male) participated in Experiment 1. Groups did not differ by age (t = -1.550, p = 0.126) but differed in sex (χ2 = 8.964, p = 0.004). 42 ADHD children (age = 9.34 ± 1.89, 27 male) participated in Experiment 2. They were naïve to the task and did not participate in Experiment 1. Participant demographic characteristics are shown in Table 1. There is currently no comparable study in ADHD that indicates the effect size for reference. Studies investigating BM perception in ASD typically have sample sizes ranged from 15 to 35 participants per group32. Considering the mild impairment of social functions in children with ADHD, we determined that a sample size of 35 - 40 participants per group would be reasonable for this study. All individuals of each group had a normal or corrected-to- normal vision and were naïve to the objectives of experiments. Written informed consent was obtained from the parents of all children before testing. The institutional review boards of the Peking University Sixth Hospital and the Institute of Psychology, Chinese Academy of Sciences have approved this study.

Revised version: Thirty-six TD children (age = 9.09 ± 2.18, 14 male) and 39 children with ADHD (age = 9.88 ± 2.23, 28 male) participated in Experiment 1. The groups did not differ by age (t = −1.550, p = 0.126) but differed in sex (χ2 = 8.964, p = 0.004). Forty-two ADHD children (age = 9.34 ± 1.89, 27 male) participated in Experiment 2. The participants were naïve to the task and did not participate in Experiment 1. The participants’ demographic characteristics are presented in Table 1. Currently, there is no comparable study on ADHD that indicates the effect size as a reference. Studies investigating BM perception in children with ASD typically have sample sizes ranging from 15 to 35 participants per group29. Considering the mild impairment of social functions in children with ADHD, we determined that a sample size of 35 to 40 participants per group was reasonable for this study. All individuals in each group had normal or corrected-to-normal vision and were naïve to the experiments objectives. Written informed consent was obtained from the parents of all the children before testing. This study was approved by the Institutional Review Boards of the Peking University Sixth Hospital and the Institute of Psychology, Chinese Academy of Sciences.

Assessment

K-SADS-PL-C DSM-5

Earlier version: K-SADS-PL-C DSM-5 is a semi-structured interview instrument for evaluating mental disorders in children and adolescents aged 6–18 years37. It involved thirty-five diagnoses based on the diagnostic criteria of DSM-5. A trained psychiatrist concluded the diagnosis by interviewing the parents and child. Chinese version showed great psychometric criterion38.

Revised version: The K-SADS-PL-C DSM-5 is a semi-structured interview instrument used to evaluate mental disorders in children and adolescents aged 6–18 years35. It involves 35 diagnoses based on the diagnostic criteria of the DSM-5. A trained psychiatrist confirmed the diagnosis by interviewing the parents and children. The Chinese version has demonstrated great psychometric properties36.

ADHD Rating Scale

Earlier version: ADHD Rating Scale is adapted from the ADHD diagnostic criteria of DSM40, which requires parents or teachers to complete the scale independently. Its Chinese version has an excellent psychometric criterion and consists of 2 subscales41: inattention (IA, nine items) and hyperactivity-impulsivity (HI, nine items). Each item is rated on a four-point Likert scale ranging from 1 (the symptom is “never or rarely”) to 4 (the symptom is “very often”). The final results will create three scores: (1) IA dimension score, (2) HI dimension score, and (3) Total score. Higher scores indicate more severe ADHD symptoms.

Revised version: The ADHD Rating Scale was adapted from the ADHD diagnostic criteria of the DSM38, which requires parents or teachers to complete the scale independently. Its Chinese version has excellent psychometric properties and consists of 2 subscales39: inattention (IA, nine items) and hyperactivity-impulsivity (HI, nine items). Each item is rated on a four-point Likert scale ranging from 1 (the symptom appears “never or rarely”) to 4 (the symptom appears “very often”). The final results create three scores: (1) IA dimension score, (2) HI dimension score, and (3) total score. Higher scores indicate more severe ADHD symptoms.

Social Responsiveness Scale

Earlier version: Social Responsiveness Scale (SRS) is a widely used quantitative measure with 65 items to assess the severity of social impairment in many mental disorders42, and the psychometric properties of the Chinese version are desirable43. It includes five subdimensions: social awareness, social cognition, social communication, social motivation, and autistic mannerisms. Each item is rated on a scale from 0 (never true) to 3 (almost always true), with higher scores indicating worse social ability.

Revised version: The Social Responsiveness Scale (SRS) is a widely used quantitative measure with 65 items used to assess the severity of social impairment in many mental disorders40, and the psychometric properties of the Chinese version are reliable41. It includes five sub-dimensions: social awareness, social cognition, social communication, social motivation, and autistic mannerisms. Each item is rated on a scale from 0 (never true) to 3 (almost always true), with higher scores indicating worse social ability.

Wechsler Intelligence Scale for Children-Fourth Edition

Earlier version: Wechsler Intelligence Scale for Children-Fourth Edition (WISC-IV) is widely used to test comprehensive intelligence in individuals aged 6–18. It contains fifteen subtests which comprise four broad areas of intellectual functioning: Verbal Comprehension, Perceptual Reasoning, Working Memory, and Processing Speed. Scores of the four broad areas constitute the full-scale intellectual quotient (FIQ).

Revised version: The Wechsler Intelligence Scale for Children-Fourth Edition (WISC-IV) is widely used to test comprehensive intelligence in individuals aged 6–18 years. It contains 15 subtests comprising four broad areas of intellectual functioning: Verbal Comprehension, Perceptual Reasoning, Working Memory, and Processing Speed. The scores in the four broad areas constitute the full-scale intellectual quotient (FIQ).

QB Test

Earlier version: QB Test is a 15-min Continuous Performance Test (CPT) for assessing inattention and impulsivity with a high-resolution infrared camera monitoring participant’s activity44. Previous psychometric studies have validated its good measurement properties45. After the test was completed, several Q scores were calculated to summarize participant performance. The Q scores were standardized based on normative data matching gender and age. The higher Q score implies more abnormal performance. In this study, we focused on the QbInattention, the Q score responsible for sustained attention, particularly when children focused on tasks.

Revised version: The QB Test is a 15-minute Continuous Performance Test (CPT) for assessing inattention and impulsivity, with a high-resolution infrared camera monitoring participant’s activity42. Previous psychometric studies have validated its good measurement properties43. After the test is completed, several Q scores are calculated to summarise the participants’ performances. The Q scores are standardised based on normative data matched for sex and age. A higher Q score implies more abnormal performance. In this study, we focused on QbInattention, the Q score indicating sustained attention, particularly when children are focused on tasks.

Stimuli and Procedure

Earlier version: Point-light BM stimuli sequences adopted in this study have been used in previous studies46, which were derived from configurations of individual walking motion on a treadmill and did not contain an overall translation. Each frame of BM sequences consisted of 13 white dots representing the human head and major joints and was displayed on a grey background (see video 1). Each walking cycle lasted 1s with 30 frames. For each trial, the initial frame of BM sequences was randomized. The whole point-lighted BM walker presented about 5.7 ° visual angle vertically. Stimuli were presented on a 14-inch monitor using MATLAB together with the PsychToolbox extensions. All subjects completed experiments in a dim-lit room with heads on the chinrest to ensure their eyes were 50 cm away from the monitor.

Revised version: Point-light BM stimuli sequences adopted in this study have been used in previous studies44, which were derived from the configurations of individual walking motions on a treadmill and did not contain an overall translation. Each frame of BM sequences consisted of 13 white dots representing the human head and major joints and was displayed on a grey background (see Video 1). Each walking cycle lasted 1s with 30 frames. For each trial, the initial frame of BM sequences was randomised. The entire point-lighted BM walker was presented at approximately a 5.7° vertical visual angle. Stimuli were presented on a 14-inch monitor, and responses were evaluated using MATLAB together with the PsychToolbox extensions. All subjects completed the experiments in a dimly lit room with their heads on a chinrest to ensure that their eyes were 50 cm away from the monitor.

Earlier version: In Experiment 1, children were required to complete three tasks that were similar to but slightly modified from the versions implemented in our previous study19 (Figure 2). Each trial began with a fixation cross (0.6° * 0.6°). Following a random interval of 1200-1500 ms, the monitor displayed a task-specific BM sequence that lasted 2 seconds (60 frames).

Revised version: In Experiment 1, the children were required to complete three tasks that were similar to but slightly modified from the versions implemented in our previous study30 (Figure 2). Each trial began with a fixation cross (0.6° * 0.6°). Following a random interval of 1200-1500 ms, the monitor displayed a task-specific BM sequence lasting for 2 s (60 frames).

Earlier version: Illustration of the trial sequence. In Task 1 (BM-local), a monitor displayed scrambled BM sequences. Participants only judged the facing direction of the scrambled walker using local BM cues. In Task 2 (BM-global), each trial only showed an intact or scrambled walker (black dots in the figure) embedded within a mask containing local BM cues. Because the two conditions contained the same local cues that were also present in the mask, a participant must rely on global BM cues to determine whether an intact walker was present in the mask. The figure showed one of five possible facing directions of an intact walker (i.e., facing participants). In Task 3 (BM-general), the stimuli sequence consisted of an intact walker (black dots) and a mask containing similar target local cues, and children judged the facing direction of the walker using general BM cues (local + global). Dots in the figure were rendered in black for better illustration and displayed in white in the actual experiments.

Revised version: Illustration of the trial sequence. In Task 1 (BM-local), a monitor displayed scrambled BM sequences. Participants only judged the facing direction of the scrambled walker using local BM cues. In Task 2 (BM-global), each trial only showed an intact or scrambled walker (black dots in the figure) embedded within a mask containing local BM cues. Because the two conditions contained the same local cues that were also present in the mask, the participant must rely on global BM cues to determine whether an intact walker was present in the mask. The figure shows one of five possible directions the intact walker could face (i.e., facing participants). In Task 3 (BM-general), the stimuli sequence consisted of an intact walker (black dots) and a mask containing similar target local cues, and children judged the direction the walker was facing using general BM cues (local + global). Dots in the figure were rendered in black for better illustration but were displayed in white in the actual experiments.

Earlier version: Task 1 (BM-Local) aimed to assess the participants’ ability to process local BM cues. During the task, the monitor only displayed a scrambled walker facing either left or right (see video 2). Specifically, 13 dots constituting the intact walker were randomly relocated within the original range of the BM walker (presented randomly within the 2D region). This manipulation disrupted the global configuration of the intact walker while retaining the local kinematics. After the display, we required children to press a left or right button to indicate the motion direction of the unidentified creature (i.e., scrambled BM walker) as accurately as possible. Children did not receive feedback regarding the accuracy of each response. A total of 30 trials were conducted, with 15 trials for each condition (left and right).

Revised version: Task 1 (BM-Local) assessed participants’ ability to process local BM cues. During the task, the monitor displayed only a scrambled walker facing either the left or right (see Video 2). Specifically, the 13 dots constituting the intact walker were randomly relocated within the original range of the BM walker (randomly presented in 2D). This manipulation disrupted the global configuration of the intact walker while retaining local kinematics. After the display, we required the children to press a left or right button to indicate the direction of motion of the unidentified creature (i.e., the scrambled BM walker) as accurately as possible. Children did not receive feedback on the accuracy of each response. Thirty trials were conducted, with 15 trials for each condition (left and right).

Earlier version: Task 2 (BM-Global) tested the ability to process the global configuration cues of the BM walker. A target walker (either scrambled or intact) was displayed within a mask (see video 3) during this task. The mask consisted of two scrambled target walkers (26 dots) with the same locomotion direction as the target walker, displayed within the boundary about 1.44 times larger than the intact walker. The scrambled or intact version of the target walker was randomly embedded into the mask and entirely overlaid by the mask. By this means, the global BM component could be isolated as two conditions (i.e., scrambled and intact walker) contained the same information of local kinematics, rendering the local motion cues uninformative. Children were required to judge whether there was an intact walker in the mask. The correct response relied on extracting global cues from the intact walker. To prevent children from learning the shape of the walker24, we set target walkers possibly facing one of five equally spaced directions from left to right. Out of the five walkers we used, two faced straight left or right, orthogonal to the viewing directions. Two walked with their bodies oriented at a 45-degree angle to the left or right of the observer. The last one walked towards the observer. Video 4 demonstrated the five facing directions. A total of 30 trials were conducted, consisting of two conditions (intact or scrambled target).

Revised version: Task 2 (BM-Global) tested the ability to process the global configuration cues of the BM walker. A target walker (scrambled or intact) was displayed within the mask (see Video 3) during this task. The mask consisted of two scrambled target walkers (26 dots) with the same locomotion direction as the target walker, displayed within a boundary approximately 1.44 times larger than the intact walker. A scrambled or intact version of the target walker was randomly embedded into the mask and entirely overlaid. Thus, the global BM component could be isolated as two conditions (i.e., scrambled and intact walkers) containing the same local kinematics information, rendering the local motion cues uninformative. The children were required to judge whether there was an intact walker on the mask. A correct response relied on the extraction of global cues from an intact walker. To prevent children from learning the shape of the walker21, we set target walkers that possibly faced one of five equally spaced directions from left to right. Of the five walkers used, two faced straight to the left or right, orthogonal to the viewing direction. Two walked with their bodies oriented at a 45-degree angle to the left or right of the observer. The last one walked towards the observer. Video 4 shows the five facing directions of the walker. Thirty trials were conducted consisting of two conditions (intact or scrambled target).

Earlier version: Task 3 (BM-General) tested the ability to process general BM cues (local + global cues). In Task 3, the monitor displayed an intact walker (facing either left or right) embedded within a mask (see video 5). The mask used in this task was similar to that in Task 2. Children were required to judge the facing direction of the target walker (left or right). Since the mask and the target walker contained the same local BM cues and the target walker is presented with additional global configuration cues, children could rely on the general BM information (i.e., a combination of local and global cues) BM information to perform the task. Task 3 consisted of 30 trials, with 15 trials for each facing directions. Other parameters of Task 2 and 3 were similar to Task 1. Before each task, children practiced five trials to make sure good understanding. We performed three tasks in a fixed order so that participants were naïve to the nature of local BM cues in Task 1.

Revised version: Task 3 (BM-General) tested participants’ ability to process general BM cues (local and global). In Task 3, the monitor displayed an intact walker (facing either the left or right) embedded within a mask (see Video 5). The mask used in this task was similar to that used in Task 2. The children were required to judge the direction the target walker was facing (left or right). Because the mask and target walker contained the same local BM cues and the target walker was presented with additional global configuration cues, children could rely on general BM information (i.e., a combination of local and global cues) to perform the task. Task 3 consisted of 30 trials, with 15 trials in each facing direction. The other parameters of Tasks 2 and 3 were similar to those of Task 1. Before each task, the children practiced for five trials to ensure a good understanding. We performed the three tasks in a fixed order so that the participants were naïve to the nature of the local BM cues in Task 1.

Earlier version: In Experiment 2, 42 children with ADHD completed the same procedures as in Experiment 1. Beyond that, their parents should fill in the Social Responsiveness Scale (SRS) to assess social interaction.

Revised version: In Experiment 2, 42 children with ADHD completed the same procedures as in Experiment 1. In addtion, the parents completed the SRS to assess social interaction.

Statistics

Earlier version: We carried out a two-sample t-test to examine the difference in BM perception abilities between TD and ADHD children47,48 and the Pearson correlation analysis was used to assess the relationship between the accuracy of each task and the SRS score. Additionally, general linear models and path analysis were used to explore the potential factors influencing BM perception. A p-value < 0.05 was considered statistically significant. Path analysis was conducted using AMOS and other analyses were conducted using SPSS.

Revised version: Two-sample t-tests were used to examine the difference in BM perception abilities between TD and ADHD children45,46, and the Pearson’s correlation analysis was used to assess the relationship between the accuracy of each task and the SRS score. Additionally, general linear models and path analyses were used to explore potential factors influencing BM perception. A p-value < 0.05 was considered statistically significant. Path analysis was conducted using AMOS, whereas other analyses were conducted using SPSS.

Results

Children with ADHD exhibit atypical BM perception

Earlier version: Figure 3 displays the mean accuracies (ACC) for both the TD and ADHD groups across the three tasks in Experiment 1. We examined the difference in ACC between the TD and ADHD groups for each task using the two-sample t-test. The results of Task 1 showed a significant difference (TD: 0.52 ± 0.13, ADHD: 0.44 ± 0.09, t73 = 3.059, p = 0.003, Cohen’s d = 0.716), indicating that children with ADHD exhibited impaired local BM processing ability. For Task 2 and 3, where children were asked to detect the presence or discriminate the facing direction of the target walker (Task 2 - TD: 0.70 ± 0.12, ADHD: 0.59 ± 0.12, t73 = 3.677, p < 0.001, Cohen’s d = 0.861; Task 3 - TD: 0.79 ± 0.12, ADHD: 0.63 ± 0.17, t73 = 4.702, p < 0.001, Cohen’s d = 1.100).

Revised version: Figure 3 displays the mean accuracies (ACC) for both the TD and ADHD groups across the three tasks in Experiment 1. We examined the difference in the ACC between the TD and ADHD groups for each task using a two-sample t-test. The results of Task 1 showed a significant difference (TD: 0.52 ± 0.13, ADHD: 0.44 ± 0.09, t73 = 3.059, p = 0.003, Cohen’s d = 0.716), indicating that children with ADHD exhibited impaired local BM processing ability. For Tasks 2 and 3, where children were asked to detect the presence or discriminate the direction the target walker was facing, the TD group had higher accuracies than the ADHD group (Task 2 - TD: 0.70 ± 0.12, ADHD: 0.59 ± 0.12, t73 = 3.677, p < 0.001, Cohen’s d = 0.861; Task 3 - TD: 0.79 ± 0.12, ADHD: 0.63 ± 0.17, t73 = 4.702, p < 0.001, Cohen’s d = 1.100). These findings suggest the presence of impaired global and general BM perception in children with ADHD. To ensure that sex did not influence the results, we conducted a subsampling analysis with balanced data, and the results remained consistent (see Supplementary Information).

Earlier version: The mean accuracy of the three tasks. Typically developing (TD) children had higher accuracies than children with attention deficit hyperactivity disorder (ADHD) in three tasks in Experiment 1. Error bars show standard deviations. ***p < 0.001

Revised version: The mean accuracy of the three tasks. Typically developing (TD) children had higher accuracies than children with attention deficit hyperactivity disorder (ADHD) in the three tasks in Experiment 1. Error bars show standard deviations. ***p < 0.001

Removed in the revised version: These findings suggest impaired global and general BM perception in children with ADHD. To further ensure that gender did not influence the results, we conducted a subsampling analysis with balanced data, and the results remained consistent (see Supplementary Information).

Earlier version: Atypical perception of local BM information predicted impaired social interaction in ADHD

Revised version: Atypical perception of local BM information predicts impaired social interaction in ADHD

Earlier version: Experiment 1 provided evidence of atypical BM perception in children with ADHD. Previous study revealed that local BM processing ability is heritable and negatively correlated with autistic traits, especially in the subdimension of communication19. In addition, substantial evidence indicates that children with ADHD often have problems in social interaction. We expected that compromised social interaction in ADHD might be associated with atypical local BM processing. To confirm this assumption, we recruited 42 naïve children with ADHD to participate in Experiment 2 and examined the relationship between their social interaction abilities and BM perception. Their parents or caregivers who took care of the children completed SRS. The SRS total score of the ADHD group was higher than that of the TD group (SRS total score - ADHD: 54.64 ± 18.42, TD: 38.64 ± 12.47, t = -4.277, p < 0.001). We found that children with higher total score of SRS performed worse in Task 1, i.e., the ability of local BM processing was negatively correlated to SRS total scores (r = - 0.264, FDR-corrected p = 0.033). We also found significant correlations between SRS total score and global BM processing as well as general BM processing (global BM perception: r = -0.238, FDR-corrected p = 0.039; general BM perception: r = -0.359, FDR-corrected p = 0.006).

Revised version: Experiment 1 provides evidence of atypical BM perception in children with ADHD. Previous studies have revealed that BM processing ability is a hallmark of social cognition10 and is negatively correlated with social ability30. Substantial evidence indicates that children with ADHD often experience problems with social interaction. We hypothesised that compromised social interaction in children with ADHD would be associated with BM processing. To confirm this hypothesis, we recruited 42 naïve children with ADHD to participate in Experiment 2 and examined the relationship between their social interaction abilities and BM perception. Parents or caregivers completed the SRS. The SRS total score of the ADHD group was higher than that of the TD group (SRS total score - ADHD: 54.64 ± 18.42, TD: 38.64 ± 12.47, t = −4.277, p < 0.001). We found that children with higher total SRS scores performed worse on the three tasks; that is, the abilities of BM processing were negatively correlated with SRS total scores (Task 1: r = −0.264, false discovery rate [FDR]-corrected p = 0.033; Task 2: r = −0.238, FDR-corrected p = 0.039; Task 3: r = −0.359, FDR-corrected p = 0.006).

Earlier version: In further subgroup analysis, as depicted in Figure 4, the correlation between the SRS total score and the ability to process local cues was only found in the ADHD group (ADHD: r = -0.461, FDR-corrected p = 0.004; TD: r = 0.109, FDR-corrected p = 0.547), particularly on subscales of social awareness, social cognition, social communication, social motivation (see Table 2 for detailed information). However, we did not find significant correlation between SRS total score and global or general BM processing in either the ADHD or TD groups separately (global BM perception - TD: r = -0.020, FDR-corrected p = 0.910, ADHD: r = -0.207, FDR-corrected p = 0.374; general BM perception - TD: r = -0.118, FDR-corrected p = 0.514, ADHD: r = -0.286, FDR-corrected p = 0.134).

Revised version: In further subgroup analysis, as depicted in Figure 4, the correlation between the SRS total score and the ability to process local cues was only found in the ADHD group (ADHD: r = −0.461, FDR-corrected p = 0.004; TD: r = 0.109, FDR-corrected p = 0.547), particularly on subscales of social awareness, social cognition, social communication, social motivation (see Table 2 for detailed information). However, we did not find statistically significant correlation between the SRS total score and global or general BM processing in either the ADHD or TD groups (global BM perception, TD: r = −0.020, FDR-corrected p = 0.910, ADHD: r = −0.207, FDR-corrected p = 0.374; general BM perception, TD: r = −0.118, FDR-corrected p = 0.514, ADHD: r = −0.286, FDR-corrected p = 0.134).

Earlier version: Correlations between response accuracies and SRS total score. The ability to process local cues is significantly correlated with the SRS total score in ADHD group. The grey shading represents a 95% confidence interval. ** FDR- corrected p < 0.01.

Revised version: Correlations between response accuracies and SRS total score. The ability to process local cues is significantly correlated with the SRS total score in the ADHD group. The grey shading represents the 95% confidence interval. ** FDR-corrected p < 0.01.

Earlier version: For an exploration of the distinct relationship between local BM processing and SRS total scores in the ADHD group, we further examined the differences among these correlations (see Supplementary Information). In Task 1, we observed a significantly stronger correlation between SRS total score and the response accuracy of the ADHD group compared to the TD group (p = 0.003). However, this distinction was not identified in Tasks 2 (p = 0.381) or Task 3 (p = 0.455). Additionally, we found that the correlation between SRS total score and the response accuracy of the ADHD group in Task 1 was slightly stronger than in Task 2 (p = 0.066) and Task 3 (p = 0.181). These findings suggested that atypical local BM processing may specifically reflect the impaired social interaction in the ADHD group.

Revised version: To explore whether the significant correlation between local BM processing and SRS total scores in the ADHD group was greater than other correlations, we constructed a general linear model to compare these correlations further (see Supplementary Information). In Task 1, we observed a significantly stronger correlation between SRS total score and response accuracy in the ADHD group than in the TD group (p = 0.003). However, this distinction was not observed for Task 2 (p = 0.381) or Task 3 (p = 0.455). Additionally, we found that the correlation between the SRS total score and response accuracy of the ADHD group in Task 1 was slightly greater than in Task 2 (p = 0.066) and Task 3 (p = 0.181). For complete information, see Supplementary Information.

Global BM processing develops with age and is regulated by reasoning intelligence and attention function

Earlier version: Experiment 2 indicated that local BM processing, as opposed to global BM processing, was more closely associated with social interaction in the ADHD group. To investigate the specificity of this association, we examined the relationship between impairments in different biological processing abilities and factors such as age, IQ, and attention function in the ADHD group. Data from the ADHD group in both Experiment 1 and Experiment 2 were integrated, resulting in a total of 80 ADHD participants (with one child not completing the QB test). Three linear models were built to investigate contributing factors: (a) BM-local = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention; (b) BM-global = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention; (c) BM-general = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention + β5 * BM-local + β6 * BM-global. We screened factors with the largest contribution to the models using step-wise regression. In model (a), no regressor has remained after step-wise regression, suggesting that the local BM processing remained stable with age and was not affected by attention and IQ. In model (b), the ability to process global BM cues was enhanced with age (standardized β1 = 0.251, p = 0.025). In model (c), higher FIQ, particularly on subdimension of Perceptual Reasoning (standardized β3 = 0.271, p = 0.005), and better performance in global BM processing (standardized β6 = 0.290, p = 0.004) predicted better performance in general BM processing. Furthermore, as children grew up, the ability to probe general BM information improved (standardized β1 = 0.365, p < 0.001). It is worth noting that QbInattention showed a strong negative correlation with Perceptual Reasoning (r = -0.355, p = 0.001) and general BM perception (r = -0.246, p = 0.028). Due to the potential issue of collinearity, we employed path analysis post hoc to visualize these relationships (Figure 5). The result indicated sustained attention (i.e., QbInattention) did not directly predict performance in BM-General (i.e., Task 3).

Revised version: Many factors can affect social cognition in ADHD, such as general cognitive impairment5, symptom severity8, or age5. Therefore, we examined the relationship between impairments in different BM processing abilities and factors such as age, IQ, and attention function in the ADHD group. Data from the ADHD group in both Experiments 1 and 2 were integrated, resulting in 80 ADHD participants (one child did not complete the QB test). Three linear models were built to investigate the contributing factors: (a) ACClocal = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention, (b) ACCglobal = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention, and (c) ACCgeneral = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * QbInattention + β5 * ACClocal + β6 * ACCglobal. ACClocal, ACCglobal and ACCgeneral refer to the response accuracies of the three tasks in the ADHD group, and QbInattention is the standardised score for sustained attention function. We screened the factors with the largest contribution to the models using step-wise regression. In model (a), no variable remained after stepwise regression, suggesting that local BM processing remained stable with age and was not affected by attention or IQ. In model (b), the ability to process global BM cues was enhanced with age (standardised β1 = 0.251, p = 0.025). In model (c), higher FIQ, particularly on the subdimension of Perceptual Reasoning (standardised β3 = 0.271, p = 0.005), and better performance in global BM processing (standardised β6 = 0.290, p = 0.004) predicted better performance in general BM processing. Furthermore, as children aged, the ability to probe general BM information improved (standardised β1 = 0.365, p < 0.001). It is worth noting that QbInattention showed a strong negative correlation with Perceptual Reasoning (r = −0.355, p = 0.001) and general BM perception (r = −0.246, p = 0.028). Owing to the potential collinearity issue, we employed a post hoc path analysis to visualise these relationships (Figure 5). The results indicated that sustained attention (i.e., QbInattention) did not directly predict performance on BM-General (i.e., Task 3) but was significantly indirectly predicted by Perceptual Reasoning ability. Furthermore, as children with ADHD aged, their performance in BM-General (i.e., Task 3) improved, both directly and through the enhanced processing of global BM cues.

Factors influencing BM perception in ADHD children. Post hoc path analysis confirmed that the effect of sustained attention on performance in BM-General was entirely mediated by Perceptual Reasoning, and the ability of global BM processing partly mediated the effect of age on performance in BM-General.

Removed in the revised version: However, a significant indirect path via the ability of Perceptual Reasoning was identified. Furthermore, as children with ADHD grew older, their performance in BM-General (i.e., Task 3) improved, both directly and through enhanced processing of global BM cues.

Earlier version: We also built three models to further explore the effects of Reasoning IQ and age on BM perception in TD children: (a) BM-local = β0 + β1 * age + β2 * gender + β3 * FIQ. (b) BM-global = β0 + β1 * age + β2 * gender + β3 * FIQ; (c) BM-general = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * BM-local + β5 * BM-global. In model (a), no regressor remained significant after the step-wise regression. However, in models (b) and (c), we observed positive relationships between age and performance (BM-global: standardized β = 0.396, p = 0.017; BM-general: standardized β = 0.330, p = 0.049).

Revised version: We also built three models to explore further the effects of Reasoning IQ and age on BM perception in TD children: (a) ACClocal = β0 + β1 * age + β2 * gender + β3 * FIQ. (b) ACCglobal = β0 + β1 * age + β2 * gender + β3 * FIQ; (c) ACCgeneral = β0 + β1 * age + β2 * gender + β3 * FIQ + β4 * ACClocal + β5 * ACCglobal. In model (a), no regressor remained significant after stepwise regression. However, in models (b) and (c), we observed positive relationships between age and performance (ACCglobal: standardised β = 0.396, p = 0.017; ACCgeneral: standardised β = 0.330, p = 0.049). We also a conducted a path analysis similar to that in the ADHD group and found no statistically significant mediator effect (Figure S3).

Earlier version: To sum up, our findings suggest that the abilities to perceive global and general BM cues, rather than local BM cues, would improve with age. Age might play different roles in the processing of different BM cues. To further verify this speculation, we examine the differences between the improvements in processing local cues and that in processing global and general cues with age (see Supplementary Information). In the ADHD group, we observed that the abilities to process general BM cues significantly improved with age compared to local cues (p < 0.001). Additionally, there was a trend indicating that the improvement in processing global BM cues with age was greater than that in processing local BM cues (p = 0. 073). However, these patterns were not observed in the TD group (BM-Local vs BM-Global: p = 0. 575; BM-Local vs BM-General: p = 0. 739).

Revised version: In summary, our findings suggest that the ability to perceive global and general BM cues, rather than local BM cues, improves with age in both groups. We speculated that age-related improvement was different when processing different BM cues and in different groups. Therefore, we examined the differences between the age-related improvement in processing local cues and that in processing global and general cues (see Supplementary Information for detail methods). In the ADHD group, we observed that the ability to process general BM cues significantly improved with age compared to local cues (p < 0.001) and a trend that the improvement in processing global BM cues with age was greater than that in processing local BM cues (p = 0. 073). However, these patterns were not observed in the TD group (see Supplementary Information). In addition, we examined the differences in the improvements in processing of BM cues with age between the two groups for each task. There was no difference in the effect of age on the response accuracy between the TD and ADHD groups for the three tasks (see Supplementary Information).

Discussion

Earlier version: Our study contributes several promising findings concerning atypical biological motion perception in ADHD. Specifically, we observe the atypical local and global BM perception in children with ADHD. Notably, a potential dissociation between the processing of local and global BM information is identified. The ability to process local BM cues appears to be linked to the traits of social interaction among children with ADHD. In contrast, global BM processing exhibits an age-related development. Additionally, general BM perception may be affected by factors including attention.

Revised version: Our study contributes several promising findings concerning atypical BM perception in children with ADHD. Specifically, we observed the atypical local and global BM perception in children with ADHD. Notably, local and global BM processing exhibited distinct features. The ability to process local BM cues appears to be associated with social interaction traits in children with ADHD. In contrast, global BM processing was associated with age-related development. In addition, general BM perception may be affected by factors such as attention.

Earlier version: BM perception is a widely studied topic in the field of visual cognition, due to its inherent biological and social properties. BM processing has significant value in successful daily life, particularly in nonverbal communication11,49 and adaptive behavior29,50. As an early emerging ability, BM can be discriminated in the newborn baby21. In typical development children, there is a clear association between BM perception and social cognitive abilities51. For example, 12-month-old infants exhibit social behaviors (i.e., following gaze) elicited by BM displays52. It is, therefore, of supreme importance of BM to children’s development of social cognition. This ‘social interpretation’ of BM proposed that the difficulties in processing BM may serve as an indicator of impaired social interaction19 This is evident in children with ASD, where impaired BM perception is a characteristic feature linked to impaired social function10,53 .

Revised version: BM perception is a widely studied topic in visual cognition owing to its inherent biological and social properties. BM processing has significant value in successfully navigating daily life, particularly in non-verbal communication11,47 and adaptive behavior26,48. In TD children, there is a clear association between BM perception and social cognitive abilities49. For example, 12-month-old infants exhibit social behaviours (i.e., following gaze) elicited by BM displays50. Therefore, BM perception plays a crucial role in the development of children’s social cognition. This ‘social interpretation’ of BM suggests that the difficulties in processing BM may serve as an indicator of impaired social interaction30. Our results are consistent with these findings. We observed atypical BM perception in children with ADHD and a significant relationship between BM perception performance and the SRS total score. Further subgroup analysis revealed a significant negative correlation between SRS total score and the accuracy of local BM processing in the ADHD group. This correlation was stronger in the ADHD group then in the TD group. The lack of a significant correlation may be due to the narrow range of SRS scores in the TD group. Future studies should increase the sample size to explore the correlation among diverse individuals. These findings suggest that BM processing is regarded as a distinct hallmark of social cognition in ADHD children10,30.

Added in the revised version: BM perception is a multi-level phenomenon51–53. At least in part, the processing of of local and global BM information appears to involve different neural mechanisms16. Sensitivity to local BM cues emerges early in life54,55 and involves rapid processing in the subcortical regions16,56–58. As a basic pre-attentive feature23, local BM cues can guide visual attention spontaneously59,60. In contrary, the ability to process global BM cues is related to slow cortical BM processing and is influenced by many factors such as attention25,26 and visual experience21,51. As mentioned above, we found a significant negative correlation between the SRS total score and the accuracy of local BM processing, specifically in the ADHD group. This could be due to decreased visual input related to atypical local BM processing, which further impairs global BM processing. According to the two-process theory of biological motion processing61, local BM cues guide visual attention towards BM stimuli55,62. Consequently, the visual input of BM stimuli increases, facilitating the development of the ability to process global BM cues through learning21,63. The latter is a prerequisite for attributing intentions to others and facilitating social interactions with other individuals20,64,65. Thus, atypical local BM processing may contribute to impaired social interaction through altered visual inputs. Further empirical studies are required to confirm these hypotheses.

Added in the revised version: The ability to process global BM cues develops with age. Previous studies indicated that global BM perception is enhanced with age in TD children66,67. This developmental trend in global BM processing is also evident in individuals with impaired BM perception (e.g., children with ASD). BM processing performance in children with ASD becomes more aligned with that of TD children as they age29,68–70. Our study contributes new evidence to the understanding of the development of global BM processing. We found that the ability to process global and general BM cues improved significantly with age in both the TD and ADHD groups, impling that the processing module for global BM cues tends to mature with development. This finding is akin to the potential age-related improvements observed in certain aspects of social cognitive deficits in individuals with ADHD5. Interestingly, in the ADHD group, the improvement in processing general and global BM cues was greater than in processing local BM cues. Few developmental studies have been conducted on local BM processing. The ability to process local BM cues remained stable and did not exhibit a learning trend21,25. A reasonable interpretation may be that local BM processing is a low-level mechanism, probably performed by the primary visual cortex and subcortical regions such as the superior colliculus, pulvinar, and ventral lateral nucleus14,56,61. Although there was no statistical difference between the improvements in processing local and global BM cues in the TD group, this may be due to the relatively small sample size or the relatively higher baseline abilities of BM perception in TD children, resulting in a relatively milder improvement. It is worth noting that the ability to process global BM cues was positively correlated with the performance in processing general BM cues in the ADHD group, whereas no such correlation was found in the TD group. This suggests that TD children are able to extract and integrate both local and global cues, whereas children with ADHD may rely more on global BM cues to judge the direction the walker is facing when presented with both local and global BM cues, which correspond to a hierarchical model51. Once a living creature is detected, an agent (i.e., is it a human?) can be recognised by a coherent, articulated body structure that is perceptually organised based on its motions (i.e., local BM cues)71. This involves top-down processing and probably requires attention25,72, particularly in the presence of competing information26. Our findings are consistent with those of previous studies on the cortical processing of BM73, as we found that the severity of inattention in children with ADHD was negatively correlated with their performance in BM processing. However, further studies are needed to verify this hypothesis.

Removed in the revised version: ADHD, a neurodevelopmental disorder, is typically identified in early childhood. Children with ADHD often show impaired social function. We observed atypical BM perception in children with ADHD, potentially contributing to social disfunctions.

Removed in the revised version: The proposition that BM perception serves as a prerequisite for attributing intentions to others, facilitating adaptive interaction with other individuals23,54,55, aligns with our discovery of a significant relationship between BM perception performance and the SRS total score. Further subgroup analysis revealed a significant negative correlation between the SRS total score and the accuracy of processing local BM specifically in the ADHD group. Notably, this correlation was stronger in the ADHD group compared to the TD group. This may be due to the narrow range of SRS scores in TD group. Future study should increase sample size to explore the correlation in diverse individuals. To sum up, these findings suggest that local BM processing could be regarded as a distinct hallmark of social cognition in ADHD children10,19.

Removed in the revised version: BM perception is considered a multi-level phenomenon56–58. At least in part, processing information of local BM and global BM appears to involve different genetic and neural mechanisms16,19. Using the classical twin method, Wang et al. found that the distinction between local and global BM processing may stem from the dissociated genetic bases. The former, to a great degree, seems to be acquired phylogenetically20,21,59,60, while the latter is primarily obtained through individual development19. The sensitivity to local rather than global BM cues seems to emerge early in life. Visually inexperienced chicks exhibit a spontaneous preference for the BM stimuli of hen, even when the configuration was scrambled20. The same finding was reported in newborns. On the contrary, the ability to process global BM cues rather than local BM cues may be influenced by attention28,29 and shaped by experience24,56. There is evidence of an age-related improvement in global BM, as a previous study indicated that global BM perception was enhanced with age in TD children61,62. This developmental trend also seems to exist in individuals with ASD. The studies in children consistently showed the impairment of global BM processing in children with ASD63,64, whereas no impairment was found in studies with regard to adults with ASD65–67. The performance of BM processing in children with ASD becomes more aligned with TD children as they get older32.

Removed in the revised version: Our study contributes new evidence to the understanding of the development of global BM processing. We found that the ability to process global and general BM cues improved significantly with age in both TD and ADHD groups, which imply the processing module for global BM cues tends to be mature with development. In the ADHD group, the improvement in processing general and global BM cues is greater than that in processing local BM cues, while no difference was found in TD group.

Removed in the revised version: This may be due to the relatively higher baseline abilities of BM perception in TD children, resulting in a relatively milder improvement. These findings also suggest a dissociation between the development of local and global BM processing. There seems to be an acquisition of ability to process global BM cues, akin to the potential age-related improvements observed in certain aspects of social cognition deficits among individuals with ADHD5, whereas local BM may be considered an intrinsic trait19. It is worth noting that the ability to process global BM cues is positively correlated with the performance of processing general BM cues in ADHD group, while no such correlation was not found in TD group. This suggests that TD children are able to extract and integrate both local and global cues, whereas children with

Removed in the revised version: ADHD may rely more on global BM cues to judge the facing direction of the walker when presented with both local and global BM cues, which correspond to a hierarchical model56. Once a living creature is detected, an agent (i.e., is it a human?) can be recognized by a coherent, articulated body structure perceptually organized from its motions (i.e., local BM cues)68. This process involves a top-down processing and probably requires attention28,69, especially in the presence of competing information29. Our findings are consistent with previous research on the cortical processing of BM70, as we found the severity of inattention in the children with ADHD was negatively correlated to their performance of BM processing. More future works are required to verify these conjectures.

Earlier version: Interestingly, children with impaired BM perception may employ a compensation strategy71. Previous study found no impairment in BM recognition in autistic individuals with high IQ71, but children with ASD exhibited weaker adaptation effects for biological motion compared to TD children72. One possibility is that individuals with high IQ and impaired BM perception are able to develop or employ reliable strategies used for BM recognition, compensating for the lack of intuitive social perceptual processing73,74. The current study supports this assumption, as children with higher IQ, particularly in Perceptual Reasoning, demonstrated better performance. Due to the impact of attention deficits on Perceptual Reasoning, the performance of children with ADHD did not align with that of TD children.

Revised version: Interestingly, children with impaired BM perception may employ a compensatory strategy74. Previous studies found no impairment in BM recognition in individuals with autism with high IQ74, but children with ASD exhibited weaker adaptation effects for biological motion than TD children75. One possibility is that individuals with a high IQ and impaired BM perception can develop or employ reliable strategies for BM recognition, compensating for the lack of intuitive social perceptual processing76,77. The current study supports this assumption, as children with a higher IQ, particularly in Perceptual Reasoning, demonstrated better performance. Owing to the impact of attention deficits on Perceptual Reasoning, the performance of children with ADHD did not align with that of TD children.

Removed in the revised version: Overall, our study first revealed two distinct and atypical fundamental abilities underlying BM perception in children with ADHD, indicating a potential dissociation of local and global BM processing. Notably, anomalous local BM processing could predict impaired social interaction in ADHD. Moreover, these results unraveled the potential contributions of age, IQ and attention to the process of BM information.

Earlier version: These findings also shed new light on the direction of future studies. Firstly, more studies are required to confirm the dissociation between the two fundamental BM processing abilities, which will contribute to understanding the respective neural mechanisms underlying the two BM processing. Secondly, exploring the performance of more advanced BM processing in children with ADHD, such as emotion and identity recognition in BM tasks, is necessary to delineate the neural profiles for the processing of biological motion in ADHD. Finally, a comparative study between ADHD and ASD is warranted to identify common neuropsychological traits and biomarkers of social cognition impairments.

Revised version: Overall, our study reveals two distinct and atypical fundamental abilities underlying BM perception in children with ADHD. Notably, anomalous local BM processing may predict impaired social interactions in children with ADHD. Moreover, these results revealed the potential contributions of age, IQ, and attention to BM information processing. These findings also shed new light for future studies. First, different developmental trends appear in local and global BM processing. Further studies are required to explore the relationship between the two fundamental BM processing abilities, which will contribute to understanding the respective and mutual neural mechanisms underlying the two types of BM processing. Second, exploring the performance of more advanced BM processing in children with ADHD, such as emotion and identity recognition in BM tasks, is necessary to delineate the neural profiles involved in processing BM in ADHD. Finally, a comparative study between ADHD and ASD is warranted to identify common neuropsychological traits and biomarkers of social cognition impairment.

Supporting information